<a href="https://colab.research.google.com/github/bfelt001/CE880_LabCourseWork/blob/main/week7/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Breanne Felton"
COLLABORATORS = "2321566"
```

---

---
# Welcome to CE880
### This is your week-7 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).

---
## Question 1:
#### Bootstrapping
Write a Python function to estimate the `mean` of the `heights` using `bootstrapping`. Your function should take the sample data as input and generate multiple bootstrap samples to compute the estimates.

<!-- grade_id: my_bootstrapping_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np

def bootstrap_estimate_height(sample_data, num_samples, sample_size):
    """
    Estimate the mean of heights using bootstrapping.

    Parameters:
    - sample_data (numpy.ndarray): Sample data containing heights.
    - num_samples (int): Number of bootstrap samples to generate.
    - sample_size (int): Size of each bootstrap sample.

    Returns:
    - mean_estimate (float): Estimated mean height.
    """
    # YOUR CODE HERE
    x = 0
    avg = []
    while x <= num_samples:
      avg.append(np.mean(np.random.choice(sample_data, sample_size, replace=False)))
      x += 1

    mean_estimate = np.mean(avg)
    return mean_estimate
    raise NotImplementedError()

bootstrap_estimate_height(sample_data, num_samples, sample_size)
```

Output:
```
175.0
```

<!-- grade_id: my_bootstrapping_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell

# Sample data (heights of students)
sample_data = np.array([160, 165, 170, 175, 180, 185, 190, 160, 165, 170, 175, 180, 185, 190])

# Parameters
num_samples = 1000
sample_size = len(sample_data)

tolerance = 0.05
assert np.allclose(bootstrap_estimate_height(sample_data, num_samples, sample_size), 175.02, rtol=tolerance)
```

---
## Question 2:

#### Majority Voting
Given a list = `[1,2,3,2,2,1,3,3,3,2,2,1]`, consider 12 judges sitting and deciding among 3 contenders. Write a function to find the winner-based on a majority voting scheme.

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
def my_majorityvote(a):
    """Write a function to find the majority voting."""
    # YOUR CODE HERE
    # Initialize counts to 0
    count1 = 0
    count2 = 0
    count3 = 0

    # Go through the votes and count each one
    for i in a:
      if i == 1:
        count1 += 1
      elif i == 2:
        count2 += 1
      else:
        count3 += 1

    # Tuple of final counts
    max_count = max(count1, count2, count3)

    # Returns majority vote
    if max_count == count1:
        return 1
    elif max_count == count2:
        return 2
    elif max_count == count3:
        return 3

    raise NotImplementedError()
```

```python
a = np.array([1,2,3,2,2,1,3,3,3,2,2,1])
my_majorityvote(a)
```

Output:
```
2
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
# Test case 1
a = np.array([1,2,3,2,2,1,3,3,3,2,2,1])
assert my_majorityvote(a) == 2

# Test case 2
b = np.array([1,2,2,2,3,3,3,3,3,1])
assert my_majorityvote(b) == 3
```

---
## Question 3:
The `Gini index` is a measure of impurity or inequality used in `decision tree` algorithms. It quantifies the likelihood of misclassifying a randomly chosen element if it were randomly classified according to the distribution of labels in the node. The formula is given as follows:


<img src="https://www.learnbymarketing.com/wp-content/uploads/2016/02/gini-index-formula.png" alt="drawing" style="width:200px;"/>

Write a function `calculate_gini_index` calculate `Gini index`. Labels is a list of class labels. Your function `calculate_gini_index` should calculates the `Gini index` based on the frequency of each class label in the dataset. Finally, it returns the Gini index value as an output.


<!-- grade_id: my_GiniIdex, points: 0, grade: false, solution: true -->
```python
def calculate_gini_index(labels):
    """
    Calculate the Gini index for a set of class labels.

    Parameters:
    - labels (list or numpy.ndarray): List of class labels.

    Returns:
    - gini_index (float): Gini index.
    """
    # Count the number of each class
    label_counts = {}
    # Cycle through and count each
    for label in labels:
        if label in label_counts:
            label_counts[label] += 1
        else:
            label_counts[label] = 1

    # Calculate total number
    total_samples = len(labels)

    # Calculate the Gini index
    gini_index = 1
    for count in label_counts.values():
        gini_index -= (count / total_samples) ** 2

    return gini_index
    raise NotImplementedError()
```

<!-- grade_id: my_GiniIdex_solution, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell

# Test case 1
labels = [1, 0, 1, 1, 1, 1, 0, 1]
gini_index = calculate_gini_index(labels)
assert gini_index == 0.375


# Test case 2
labels = [1, 0, 0, 0, 1, 1, 0, 1]
gini_index = calculate_gini_index(labels)
assert gini_index == 0.5

```