<a href="https://colab.research.google.com/github/bfelt001/CE880_LabCourseWork/blob/main/week6/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Breanne Felton"
COLLABORATORS = "2321566"
```

---

---
# Welcome to CE880
### This is your week-6 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).

---
## Question 1:

Write a Python function to calculate the `false negative rate`, you need the `true labels` as well. The `false negative rate` is the proportion of `actual positive cases` that were incorrectly classified as `negative`.

<!-- grade_id: my_ttest_ans, points: 0, grade: false, solution: true -->
```python
def calculate_false_negative_rate(y_true, y_pred):
    """
    Calculate the false negative rate given the true labels and predicted labels.

    Parameters:
    - y_true (numpy.ndarray or list): True labels.
    - y_pred (numpy.ndarray or list): Predicted labels.

    Returns:
    - false_negative_rate (float): False negative rate.
    """
    # YOUR CODE HERE
    tp = 0
    tn = 0
    fp = 0
    fn = 0
    for x, y in zip(y_true, y_pred):
        if x == 1 and y == 1:
          tp += 1
        elif x==1 and y == 0:
          fn += 1
        elif x==0 and y == 0:
          tn += 1
        else:
          fp += 1

    false_negative_rate = fn / (tp + fn)
    return false_negative_rate
    raise NotImplementedError()
```

<!-- grade_id: my_ttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
# Case 1
y_true = [1, 0, 1, 1, 0, 0, 1, 0, 1]
y_pred = [1, 0, 0, 1, 0, 1, 1, 0, 1]
assert calculate_false_negative_rate(y_true, y_pred) == 0.2

# Case 2
y_true = [1, 0, 1, 1, 0, 1, 0, 1]
y_pred = [1, 1, 0, 1, 0, 0, 1, 1]
assert calculate_false_negative_rate(y_true, y_pred) == 0.4
```

---
## Question 2:

Write Python code to calculate the `F1-score`, you need the `true labels` as well as the `predicted labels`. The `F1-score` is a measure of a model's accuracy that considers both the `precision` and `recall` of the model.

<img src="https://lh6.googleusercontent.com/JJuhZfvr8tq3asFRvSJwb1CwqMLKCfZGpjy6Cng35RSvq-5eEa9I9TPUjs1M6jgTSP5H-EiXrr8URnoiRICHWPGDE-fwMKmZNwVDBARhO4LcUgpO4Nj_bfLLNCshBc3vrYgqUO2Nj-THYwkmbr1J9fE" alt="drawing" style="width:400px;"/>

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
def calculate_f1_score(y_true, y_pred):
    """
    Calculate the F1-score given the true labels and predicted labels.

    Parameters:
    - y_true (numpy.ndarray or list): True labels.
    - y_pred (numpy.ndarray or list): Predicted labels.

    Returns:
    - f1_score (float): F1-score.
    """
    # YOUR CODE HERE
    tp = 0
    tn = 0
    fp = 0
    fn = 0
    for x, y in zip(y_true, y_pred):
        if x ==  1 and y ==  1:
            tp +=  1
        elif x ==  1 and y ==  0:
            fn +=  1
        elif x ==  0 and y ==  1:
            fp +=  1
        else:
            tn +=  1

    precision = tp / (tp + fp)
    recall = tp / (tp + fn)
    f1_score = round(2 * ((precision * recall)/(precision + recall)),2)
    print(f1_score)
    return f1_score

    raise NotImplementedError()

```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
# Case 1 data
y_true = [1, 0, 1, 1, 0, 1, 0, 1]
y_pred = [1, 1, 0, 1, 0, 0, 1, 1]

assert calculate_f1_score(y_true, y_pred) == 0.6


# # Case 2 data
y_true = [1, 0, 1, 1, 0, 1, 0, 1, 0, 1]
y_pred = [1, 1, 0, 1, 0, 0, 1, 1, 1, 1]

assert calculate_f1_score(y_true, y_pred) == 0.62

```

Output:
```
0.6
0.62
```

```python
from sklearn.metrics import fbeta_score

y_true = [1,  0,  1,  1,  0,  1,  0,  1]
y_pred = [1,  1,  0,  1,  0,  0,  1,  1]

# Calculate F-beta score with beta=1
f_beta = fbeta_score(y_true, y_pred, beta=1)

print('F-beta Score:', f_beta)

```

Output:
```
F-beta Score: 0.6
```

---
## Question 3:

#### Shapiro-Wilk Test
Tests whether a data sample has a Gaussian distribution.

Assumptions

* Observations in each sample are independent and identically distributed (iid).

Interpretation

* H0: the sample has a Gaussian distribution.
* H1: the sample does not have a Gaussian distribution.

<!-- grade_id: my_normalityttest_ans, points: 0, grade: false, solution: true -->
```python
from scipy.stats import shapiro
data = [0.873, 2.817, 0.121, -0.945, -0.055, -1.436, 0.360, -1.478, -1.637, -1.869]

def my_normalityttest(data):
    """Write code to to check the normality of the data. In other words,
    you are checking if the data is normally distributed and return p value.
    Hint: Please see what is shapiro test and how you can use it
    https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.shapiro.html"""
    # YOUR CODE HERE
    normal = shapiro(data)
    return normal.pvalue

    raise NotImplementedError()

my_normalityttest(data)
```

Output:
```
0.1934097558259964
```

<!-- grade_id: my_normalityttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_normalityttest(data), 0.193, rel_tol = 0.05)
```

```python

```